# No SInE strategy applied
# Auto-Mode selected heuristic G_E___107_B03_F1_PI_AE_Q4_CS_SP_PS_S0Y
# and selection function SelectMaxLComplexAvoidPosPred.
#
# Preprocessing time       : 0.010 s
# Presaturation interreduction done

# No proof found!
# SZS status CounterSatisfiable
# SZS output start Saturation
fof(t146_tmap_1, conjecture, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v1_borsuk_1(X3,X1)&m1_pre_topc(X3,X1)))=>![X4]:((~(v2_struct_0(X4))&(v1_borsuk_1(X4,X1)&m1_pre_topc(X4,X1)))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t146_tmap_1)).
fof(dt_k10_tmap_1, axiom, (![X1]:![X2]:![X3]:![X4]:![X5]:![X6]:(((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))&((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))&((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))&((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))&((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))&(v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))))))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))), file('atp/t146_tmap_1__8_premises.E_input', dt_k10_tmap_1)).
fof(d13_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6)))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))=>(X7=k10_tmap_1(X1,X2,X3,X4,X5,X6)<=>(r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)&r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6)))))))))))), file('atp/t146_tmap_1__8_premises.E_input', d13_tmap_1)).
fof(t141_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))=>((m1_pre_topc(X3,X4)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6))&((r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6)=>m1_pre_topc(X3,X4))&((m1_pre_topc(X4,X3)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5))&(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5)=>m1_pre_topc(X4,X3))))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t141_tmap_1)).
fof(t143_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((~(v2_struct_0(X5))&m1_pre_topc(X5,X1))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X7]:((v1_funct_1(X7)&(v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>![X8]:((v1_funct_1(X8)&(v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))&m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))))=>~(((r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X6),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X7)))&((r1_tsep_1(X3,X5)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X5),X6),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X3,X5),X8)))&((r1_tsep_1(X4,X5)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X4,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X4,X5),X7),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X4,X5),X8)))&~(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8))))))))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t143_tmap_1)).
fof(t140_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X5)&r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X6))<=>r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t140_tmap_1)).
fof(t142_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>((r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6)))=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X4,X3)),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),k10_tmap_1(X1,X2,X4,X3,X6,X5)))))))))), file('atp/t146_tmap_1__8_premises.E_input', t142_tmap_1)).
fof(t145_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(r1_tsep_1(X3,X4)=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>(r4_tsep_1(X1,X3,X4)=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t145_tmap_1)).
fof(t144_tmap_1, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>((r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))&r4_tsep_1(X1,X3,X4))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2)))))))))))))))), file('atp/t146_tmap_1__8_premises.E_input', t144_tmap_1)).
fof(c_0_9, negated_conjecture, (~(![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&(v1_borsuk_1(X3,X1)&m1_pre_topc(X3,X1)))=>![X4]:((~(v2_struct_0(X4))&(v1_borsuk_1(X4,X1)&m1_pre_topc(X4,X1)))=>(~(r1_tsep_1(X3,X4))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&(v5_pre_topc(X5,X3,X2)&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2)))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&(v5_pre_topc(X6,X4,X2)&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2)))))))=>(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))=>(v1_funct_1(k10_tmap_1(X1,X2,X3,X4,X5,X6))&(v1_funct_2(k10_tmap_1(X1,X2,X3,X4,X5,X6),u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))&(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)&m1_subset_1(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))))))))))))))), inference(assume_negation,[status(cth)],[t146_tmap_1])).
fof(c_0_10, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((v1_funct_1(k10_tmap_1(X7,X8,X9,X10,X11,X12))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|((v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))|((v2_struct_0(X9)|~m1_pre_topc(X9,X7))|((v2_struct_0(X10)|~m1_pre_topc(X10,X7))|((~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))))))&((v1_funct_2(k10_tmap_1(X7,X8,X9,X10,X11,X12),u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|((v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))|((v2_struct_0(X9)|~m1_pre_topc(X9,X7))|((v2_struct_0(X10)|~m1_pre_topc(X10,X7))|((~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))))))&(m1_subset_1(k10_tmap_1(X7,X8,X9,X10,X11,X12),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))))|((v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))|((v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))|((v2_struct_0(X9)|~m1_pre_topc(X9,X7))|((v2_struct_0(X10)|~m1_pre_topc(X10,X7))|((~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))))))))), inference(distribute,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[dt_k10_tmap_1])])])])).
fof(c_0_11, negated_conjecture, (((~v2_struct_0(esk1_0)&(v2_pre_topc(esk1_0)&l1_pre_topc(esk1_0)))&((~v2_struct_0(esk2_0)&(v2_pre_topc(esk2_0)&l1_pre_topc(esk2_0)))&((~v2_struct_0(esk3_0)&(v1_borsuk_1(esk3_0,esk1_0)&m1_pre_topc(esk3_0,esk1_0)))&((~v2_struct_0(esk4_0)&(v1_borsuk_1(esk4_0,esk1_0)&m1_pre_topc(esk4_0,esk1_0)))&(~r1_tsep_1(esk3_0,esk4_0)&((v1_funct_1(esk5_0)&(v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk5_0,esk3_0,esk2_0)&m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0)))))))&((v1_funct_1(esk6_0)&(v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))&(v5_pre_topc(esk6_0,esk4_0,esk2_0)&m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0)))))))&(r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0))&(~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|(~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|(~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0)))))))))))))))), inference(skolemize,[status(esa)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_9])])])])])])).
cnf(c_0_12,plain,(v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X6)|m1_subset_1(k10_tmap_1(X6,X3,X5,X2,X4,X1),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X6,X5,X2)),u1_struct_0(X3))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(X3))|~v1_funct_1(X4)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X5,X6)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X6)|~v2_pre_topc(X6)), inference(split_conjunct,[status(thm)],[c_0_10]), ['final']).
cnf(c_0_13,negated_conjecture,(v1_funct_1(esk5_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_14,negated_conjecture,(v1_funct_1(esk6_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_15,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_13]), ['final']).
cnf(c_0_16,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_14]), ['final']).
cnf(c_0_17,plain,(v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X6)|v1_funct_1(k10_tmap_1(X6,X3,X5,X2,X4,X1))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(X3))|~v1_funct_1(X4)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X5,X6)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X6)|~v2_pre_topc(X6)), inference(split_conjunct,[status(thm)],[c_0_10]), ['final']).
cnf(c_0_18,negated_conjecture,(m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_19,negated_conjecture,(v1_funct_2(esk5_0,u1_struct_0(esk3_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_20,negated_conjecture,(l1_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_21,negated_conjecture,(v2_pre_topc(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_22,negated_conjecture,(~v2_struct_0(esk3_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_23,negated_conjecture,(~v2_struct_0(esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_24,negated_conjecture,(m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_25,negated_conjecture,(v1_funct_2(esk6_0,u1_struct_0(esk4_0),u1_struct_0(esk2_0))), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_26,negated_conjecture,(~v2_struct_0(esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_27,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_13]), ['final']).
cnf(c_0_28,plain,(v2_struct_0(X2)|v2_struct_0(X5)|v2_struct_0(X3)|v2_struct_0(X6)|v1_funct_2(k10_tmap_1(X6,X3,X5,X2,X4,X1),u1_struct_0(k1_tsep_1(X6,X5,X2)),u1_struct_0(X3))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(X3))))|~v1_funct_2(X1,u1_struct_0(X2),u1_struct_0(X3))|~v1_funct_1(X1)|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X3))))|~v1_funct_2(X4,u1_struct_0(X5),u1_struct_0(X3))|~v1_funct_1(X4)|~m1_pre_topc(X2,X6)|~m1_pre_topc(X5,X6)|~l1_pre_topc(X3)|~v2_pre_topc(X3)|~l1_pre_topc(X6)|~v2_pre_topc(X6)), inference(split_conjunct,[status(thm)],[c_0_10]), ['final']).
cnf(c_0_29,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_14]), ['final']).
cnf(c_0_30,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_14]), ['final']).
cnf(c_0_31,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_13]), ['final']).
cnf(c_0_32,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,X2,esk5_0,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_18]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_33,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0,X3))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_24]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_34,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,X2,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,X2)),u1_struct_0(esk2_0))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_27, c_0_18]), c_0_19]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_35,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,X2,esk5_0,X3),u1_struct_0(k1_tsep_1(X1,esk3_0,X2)),u1_struct_0(esk2_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_18]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_36,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,X2)),u1_struct_0(esk2_0))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_29, c_0_24]), c_0_25]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_37,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0,X3),u1_struct_0(k1_tsep_1(X1,esk4_0,X2)),u1_struct_0(esk2_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_24]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
fof(c_0_38, plain, (![X1]:![X2]:![X3]:(epred1_3(X3,X2,X1)<=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))=>((m1_pre_topc(X3,X4)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6))&((r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6)=>m1_pre_topc(X3,X4))&((m1_pre_topc(X4,X3)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5))&(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5)=>m1_pre_topc(X4,X3))))))))))), introduced(definition)).
cnf(c_0_39,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,X2,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,X2)),u1_struct_0(esk2_0))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_30, c_0_18]), c_0_19]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_40,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,X2)),u1_struct_0(esk2_0))))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_31, c_0_24]), c_0_25]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_41,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_42,negated_conjecture,(m1_pre_topc(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_43,negated_conjecture,(m1_pre_topc(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_44,negated_conjecture,(l1_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_45,negated_conjecture,(v2_pre_topc(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_46,negated_conjecture,(~v2_struct_0(esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_47,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_48,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_49,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
fof(c_0_50, plain, (![X8]:![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:((((((((((((r2_funct_2(u1_struct_0(X10),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X10,X14),X12)|X14!=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r1_tsep_1(X10,X11))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))&(((((((((r2_funct_2(u1_struct_0(X11),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X11,X14),X13)|X14!=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r1_tsep_1(X10,X11))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))))&((((((((((~r2_funct_2(u1_struct_0(X10),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X10,X14),X12)|~r2_funct_2(u1_struct_0(X11),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X11,X14),X13))|X14=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r1_tsep_1(X10,X11))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))))&(((((((((((r2_funct_2(u1_struct_0(X10),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X10,X14),X12)|X14!=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X8,X10,X11)),u1_struct_0(X9),k3_tmap_1(X8,X9,X10,k2_tsep_1(X8,X10,X11),X12),k3_tmap_1(X8,X9,X11,k2_tsep_1(X8,X10,X11),X13)))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))&(((((((((r2_funct_2(u1_struct_0(X11),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X11,X14),X13)|X14!=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X8,X10,X11)),u1_struct_0(X9),k3_tmap_1(X8,X9,X10,k2_tsep_1(X8,X10,X11),X12),k3_tmap_1(X8,X9,X11,k2_tsep_1(X8,X10,X11),X13)))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8)))))&((((((((((~r2_funct_2(u1_struct_0(X10),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X10,X14),X12)|~r2_funct_2(u1_struct_0(X11),u1_struct_0(X9),k3_tmap_1(X8,X9,k1_tsep_1(X8,X10,X11),X11,X14),X13))|X14=k10_tmap_1(X8,X9,X10,X11,X12,X13))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X8,X10,X11)),u1_struct_0(X9)))))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X8,X10,X11)),u1_struct_0(X9),k3_tmap_1(X8,X9,X10,k2_tsep_1(X8,X10,X11),X12),k3_tmap_1(X8,X9,X11,k2_tsep_1(X8,X10,X11),X13)))|(~v1_funct_1(X13)|(~v1_funct_2(X13,u1_struct_0(X11),u1_struct_0(X9))|~m1_subset_1(X13,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X9)))))))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X9))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X9)))))))|(v2_struct_0(X11)|~m1_pre_topc(X11,X8)))|(v2_struct_0(X10)|~m1_pre_topc(X10,X8)))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[d13_tmap_1])])])])])])])).
cnf(c_0_51,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_34, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_52,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(k1_tsep_1(X1,esk3_0,esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_53,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_36, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_54,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(k1_tsep_1(X1,esk4_0,esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
fof(c_0_55, axiom, (![X1]:((~(v2_struct_0(X1))&(v2_pre_topc(X1)&l1_pre_topc(X1)))=>![X2]:((~(v2_struct_0(X2))&(v2_pre_topc(X2)&l1_pre_topc(X2)))=>![X3]:((~(v2_struct_0(X3))&m1_pre_topc(X3,X1))=>epred1_3(X3,X2,X1))))), inference(apply_def,[status(thm)],[t141_tmap_1, c_0_38])).
fof(c_0_56, plain, (![X9]:![X10]:![X11]:![X12]:![X13]:![X14]:![X15]:![X16]:((((((((((((((~r1_tsep_1(X12,X13)|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r1_tsep_1(X11,X13))|~r1_tsep_1(X11,X12))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(((((((((((~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X12,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X12,X13),X15),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X12,X13),X16))|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r1_tsep_1(X11,X13))|~r1_tsep_1(X11,X12))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))))&((((((((((((~r1_tsep_1(X12,X13)|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X13),X14),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X11,X13),X16)))|~r1_tsep_1(X11,X12))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(((((((((((~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X12,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X12,X13),X15),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X12,X13),X16))|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X13),X14),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X11,X13),X16)))|~r1_tsep_1(X11,X12))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))))&(((((((((((((~r1_tsep_1(X12,X13)|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r1_tsep_1(X11,X13))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X12)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X12),X14),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X11,X12),X15)))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(((((((((((~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X12,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X12,X13),X15),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X12,X13),X16))|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r1_tsep_1(X11,X13))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X12)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X12),X14),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X11,X12),X15)))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))))&((((((((((((~r1_tsep_1(X12,X13)|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X13),X14),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X11,X13),X16)))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X12)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X12),X14),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X11,X12),X15)))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9))))&(((((((((((~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X12,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X12,X13),X15),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X12,X13),X16))|r1_funct_2(u1_struct_0(k1_tsep_1(X9,k1_tsep_1(X9,X11,X12),X13)),u1_struct_0(X10),u1_struct_0(k1_tsep_1(X9,X11,k1_tsep_1(X9,X12,X13))),u1_struct_0(X10),k10_tmap_1(X9,X10,k1_tsep_1(X9,X11,X12),X13,k10_tmap_1(X9,X10,X11,X12,X14,X15),X16),k10_tmap_1(X9,X10,X11,k1_tsep_1(X9,X12,X13),X14,k10_tmap_1(X9,X10,X12,X13,X15,X16))))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X13)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X13),X14),k3_tmap_1(X9,X10,X13,k2_tsep_1(X9,X11,X13),X16)))|~r2_funct_2(u1_struct_0(k2_tsep_1(X9,X11,X12)),u1_struct_0(X10),k3_tmap_1(X9,X10,X11,k2_tsep_1(X9,X11,X12),X14),k3_tmap_1(X9,X10,X12,k2_tsep_1(X9,X11,X12),X15)))|(~v1_funct_1(X16)|(~v1_funct_2(X16,u1_struct_0(X13),u1_struct_0(X10))|~m1_subset_1(X16,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X13),u1_struct_0(X10)))))))|(~v1_funct_1(X15)|(~v1_funct_2(X15,u1_struct_0(X12),u1_struct_0(X10))|~m1_subset_1(X15,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X12),u1_struct_0(X10)))))))|(~v1_funct_1(X14)|(~v1_funct_2(X14,u1_struct_0(X11),u1_struct_0(X10))|~m1_subset_1(X14,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X11),u1_struct_0(X10)))))))|(v2_struct_0(X13)|~m1_pre_topc(X13,X9)))|(v2_struct_0(X12)|~m1_pre_topc(X12,X9)))|(v2_struct_0(X11)|~m1_pre_topc(X11,X9)))|(v2_struct_0(X10)|(~v2_pre_topc(X10)|~l1_pre_topc(X10))))|(v2_struct_0(X9)|(~v2_pre_topc(X9)|~l1_pre_topc(X9)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t143_tmap_1])])])])])])])).
cnf(c_0_57,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_39, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_58,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(X1,esk3_0,esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_59,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_40, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_60,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(k1_tsep_1(X1,esk4_0,esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_61,negated_conjecture,(v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_41, c_0_42]), c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_62,negated_conjecture,(v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_47, c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_63,negated_conjecture,(v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_48, c_0_43]), c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_64,negated_conjecture,(v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_49, c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_65,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r1_tsep_1(X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|X7!=k10_tmap_1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_66,negated_conjecture,(m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_51, c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_67,negated_conjecture,(v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_52, c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_68,negated_conjecture,(m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_53, c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_69,negated_conjecture,(v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_54, c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
fof(c_0_70, plain, (![X4]:![X5]:![X6]:((v2_struct_0(X4)|(~v2_pre_topc(X4)|~l1_pre_topc(X4)))|((v2_struct_0(X5)|(~v2_pre_topc(X5)|~l1_pre_topc(X5)))|((v2_struct_0(X6)|~m1_pre_topc(X6,X4))|epred1_3(X6,X5,X4))))), inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_55])])])])])])).
cnf(c_0_71,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r1_tsep_1(X3,X4)|~r1_tsep_1(X3,X5)|~r1_tsep_1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_72,negated_conjecture,(m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_57, c_0_42]), c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_73,negated_conjecture,(v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_58, c_0_42]), c_0_43]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_74,negated_conjecture,(m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0))))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_59, c_0_43]), c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_75,negated_conjecture,(v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_60, c_0_43]), c_0_42]), c_0_44]), c_0_45])]), c_0_46]), ['final']).
fof(c_0_76, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((((((((~r2_funct_2(u1_struct_0(X9),u1_struct_0(X8),k3_tmap_1(X7,X8,k1_tsep_1(X7,X9,X10),X9,k10_tmap_1(X7,X8,X9,X10,X11,X12)),X11)|~r2_funct_2(u1_struct_0(X10),u1_struct_0(X8),k3_tmap_1(X7,X8,k1_tsep_1(X7,X9,X10),X10,k10_tmap_1(X7,X8,X9,X10,X11,X12)),X12))|r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((((r2_funct_2(u1_struct_0(X9),u1_struct_0(X8),k3_tmap_1(X7,X8,k1_tsep_1(X7,X9,X10),X9,k10_tmap_1(X7,X8,X9,X10,X11,X12)),X11)|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((((r2_funct_2(u1_struct_0(X10),u1_struct_0(X8),k3_tmap_1(X7,X8,k1_tsep_1(X7,X9,X10),X10,k10_tmap_1(X7,X8,X9,X10,X11,X12)),X12)|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t140_tmap_1])])])])])])])).
cnf(c_0_77,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_61]), ['final']).
cnf(c_0_78,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_62]), ['final']).
cnf(c_0_79,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_63]), ['final']).
fof(c_0_80, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:((((((((~r1_tsep_1(X9,X10)|r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(k1_tsep_1(X7,X10,X9)),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),k10_tmap_1(X7,X8,X10,X9,X12,X11)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12))|r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(k1_tsep_1(X7,X10,X9)),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),k10_tmap_1(X7,X8,X10,X9,X12,X11)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t142_tmap_1])])])])])])])).
cnf(c_0_81,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),X5),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X4)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_12, c_0_64]), ['final']).
fof(c_0_82, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((((((((v1_funct_1(k10_tmap_1(X7,X8,X9,X10,X11,X12))|~r4_tsep_1(X7,X9,X10))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|~r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((((v1_funct_2(k10_tmap_1(X7,X8,X9,X10,X11,X12),u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))|~r4_tsep_1(X7,X9,X10))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|~r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((((v5_pre_topc(k10_tmap_1(X7,X8,X9,X10,X11,X12),k1_tsep_1(X7,X9,X10),X8)|~r4_tsep_1(X7,X9,X10))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|~r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((((m1_subset_1(k10_tmap_1(X7,X8,X9,X10,X11,X12),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))))|~r4_tsep_1(X7,X9,X10))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|~r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t145_tmap_1])])])])])])])).
cnf(c_0_83,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r1_tsep_1(esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_66]), c_0_67]), c_0_62]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_84,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r1_tsep_1(X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|X7!=k10_tmap_1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_85,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r1_tsep_1(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_68]), c_0_69]), c_0_64]), c_0_43]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_23]), c_0_46]), ['final']).
fof(c_0_86, plain, (![X1]:![X2]:![X3]:(epred1_3(X3,X2,X1)=>![X4]:((~(v2_struct_0(X4))&m1_pre_topc(X4,X1))=>![X5]:((v1_funct_1(X5)&(v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))&m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))))=>![X6]:((v1_funct_1(X6)&(v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))&m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))))=>(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))=>((m1_pre_topc(X3,X4)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6))&((r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X6)=>m1_pre_topc(X3,X4))&((m1_pre_topc(X4,X3)=>r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5))&(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(X3),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),X5)=>m1_pre_topc(X4,X3))))))))))), inference(split_equiv,[status(thm)],[c_0_38])).
cnf(c_0_87,plain,(epred1_3(X1,X2,X3)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_pre_topc(X1,X3)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~l1_pre_topc(X3)|~v2_pre_topc(X3)), inference(split_conjunct,[status(thm)],[c_0_70]), ['final']).
fof(c_0_88, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((((((((v1_funct_1(k10_tmap_1(X7,X8,X9,X10,X11,X12))|(~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12))|~r4_tsep_1(X7,X9,X10)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((((v1_funct_2(k10_tmap_1(X7,X8,X9,X10,X11,X12),u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))|(~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12))|~r4_tsep_1(X7,X9,X10)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&(((((((((v5_pre_topc(k10_tmap_1(X7,X8,X9,X10,X11,X12),k1_tsep_1(X7,X9,X10),X8)|(~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12))|~r4_tsep_1(X7,X9,X10)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7))))&((((((((m1_subset_1(k10_tmap_1(X7,X8,X9,X10,X11,X12),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8))))|(~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12))|~r4_tsep_1(X7,X9,X10)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|(~v5_pre_topc(X12,X10,X8)|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8))))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|(~v5_pre_topc(X11,X9,X8)|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8))))))))|r1_tsep_1(X9,X10))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|(v2_struct_0(X9)|~m1_pre_topc(X9,X7)))|(v2_struct_0(X8)|(~v2_pre_topc(X8)|~l1_pre_topc(X8))))|(v2_struct_0(X7)|(~v2_pre_topc(X7)|~l1_pre_topc(X7)))))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[t144_tmap_1])])])])])])])).
cnf(c_0_89,negated_conjecture,(~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))), inference(split_conjunct,[status(thm)],[c_0_11])).
cnf(c_0_90,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk3_0,esk3_0)),X4,k10_tmap_1(X1,esk2_0,X3,k1_tsep_1(esk1_0,esk3_0,esk3_0),X5,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_91,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk3_0,esk4_0)),X4,k10_tmap_1(X1,esk2_0,X3,k1_tsep_1(esk1_0,esk3_0,esk4_0),X5,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_72]), c_0_73]), c_0_61]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_92,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk4_0,esk3_0)),X4,k10_tmap_1(X1,esk2_0,X3,k1_tsep_1(esk1_0,esk4_0,esk3_0),X5,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_93,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,k1_tsep_1(esk1_0,esk4_0,esk4_0)),X4,k10_tmap_1(X1,esk2_0,X3,k1_tsep_1(esk1_0,esk4_0,esk4_0),X5,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_94,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X6)|~r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X5)), inference(split_conjunct,[status(thm)],[c_0_76]), ['final']).
cnf(c_0_95,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),esk3_0,k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,esk3_0),X4,k10_tmap_1(X1,esk2_0,X3,esk3_0,X5,esk5_0)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,esk3_0)|~r1_tsep_1(X2,esk3_0)|~r1_tsep_1(X2,X3)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_18]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_96,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,X3),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,X3,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,X3),esk4_0,k10_tmap_1(X1,esk2_0,X2,X3,X4,X5),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,X3,esk4_0),X4,k10_tmap_1(X1,esk2_0,X3,esk4_0,X5,esk6_0)))|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(esk2_0))))|~m1_subset_1(X4,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(esk2_0))|~v1_funct_2(X4,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X5)|~v1_funct_1(X4)|~r1_tsep_1(X3,esk4_0)|~r1_tsep_1(X2,esk4_0)|~r1_tsep_1(X2,X3)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_71, c_0_24]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_97,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r1_tsep_1(X3,X4)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X5),X6),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X3,X5),X8))|~r1_tsep_1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_98,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0))), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_99,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_62]), ['final']).
cnf(c_0_100,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_63]), ['final']).
cnf(c_0_101,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_64]), ['final']).
cnf(c_0_102,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_61]), ['final']).
cnf(c_0_103,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_62]), ['final']).
cnf(c_0_104,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_63]), ['final']).
cnf(c_0_105,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_64]), ['final']).
cnf(c_0_106,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_61]), ['final']).
cnf(c_0_107,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_62]), ['final']).
cnf(c_0_108,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_63]), ['final']).
cnf(c_0_109,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r1_tsep_1(X3,X4)|~r1_tsep_1(X3,X5)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X4,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X4,X5),X7),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X4,X5),X8))), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_110,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_64]), ['final']).
cnf(c_0_111,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_61]), ['final']).
cnf(c_0_112,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X4,X3)),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),k10_tmap_1(X1,X2,X4,X3,X6,X5))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r1_tsep_1(X3,X4)), inference(split_conjunct,[status(thm)],[c_0_80]), ['final']).
cnf(c_0_113,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X6),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X7))|~r1_tsep_1(X3,X5)|~r1_tsep_1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_114,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_62]), ['final']).
cnf(c_0_115,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_63]), ['final']).
cnf(c_0_116,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_64]), ['final']).
cnf(c_0_117,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~r1_tsep_1(X3,X4)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v5_pre_topc(X5,X3,X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v5_pre_topc(X6,X4,X2)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r4_tsep_1(X1,X3,X4)), inference(split_conjunct,[status(thm)],[c_0_82]), ['final']).
cnf(c_0_118,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_61]), ['final']).
cnf(c_0_119,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_13]), ['final']).
cnf(c_0_120,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_77, c_0_14]), ['final']).
cnf(c_0_121,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_13]), ['final']).
cnf(c_0_122,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_78, c_0_14]), ['final']).
cnf(c_0_123,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_13]), ['final']).
cnf(c_0_124,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_79, c_0_14]), ['final']).
cnf(c_0_125,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_13]), ['final']).
cnf(c_0_126,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X4),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_81, c_0_14]), ['final']).
cnf(c_0_127,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_62]), ['final']).
cnf(c_0_128,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_63]), ['final']).
cnf(c_0_129,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_64]), ['final']).
cnf(c_0_130,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk5_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_15, c_0_61]), ['final']).
cnf(c_0_131,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_62]), ['final']).
cnf(c_0_132,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_63]), ['final']).
cnf(c_0_133,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X6),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X7))|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X5),X6),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X3,X5),X8))|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X4,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X4,X5),X7),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X4,X5),X8))), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_134,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_64]), ['final']).
cnf(c_0_135,negated_conjecture,(v5_pre_topc(esk5_0,esk3_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_136,negated_conjecture,(v5_pre_topc(esk6_0,esk4_0,esk2_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_137,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,X2,X3,X4,esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|v2_struct_0(X1)|v2_struct_0(X3)|v2_struct_0(X2)|v2_struct_0(X4)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(esk6_0,u1_struct_0(X3),u1_struct_0(X2))|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~l1_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X1)|~v2_pre_topc(X2)), inference(spm,[status(thm)],[c_0_16, c_0_61]), ['final']).
cnf(c_0_138,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0)|k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,X1,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_83, c_0_18]), c_0_19]), c_0_13])])).
cnf(c_0_139,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r1_tsep_1(X3,X4)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X5),X6),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X3,X5),X8))|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X4,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X4,X5),X7),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X4,X5),X8))), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_140,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X6),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X7))|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X5),X6),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X3,X5),X8))|~r1_tsep_1(X4,X5)), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_141,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_74]), c_0_75]), c_0_63]), c_0_42]), c_0_43]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_22]), c_0_26]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_142,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,X2,X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~v1_funct_1(X2)|~r1_tsep_1(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_65, c_0_74]), c_0_75]), c_0_63]), c_0_42]), c_0_43]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_22]), c_0_26]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_143,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,X1,esk6_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_85, c_0_24]), c_0_25]), c_0_14])])).
fof(c_0_144, plain, (![X7]:![X8]:![X9]:![X10]:![X11]:![X12]:(((((((~m1_pre_topc(X9,X10)|r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(X10),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),X12))|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|~epred1_3(X9,X8,X7))&(((((((~r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(X10),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),X12)|m1_pre_topc(X9,X10))|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|~epred1_3(X9,X8,X7))&(((((((~m1_pre_topc(X10,X9)|r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(X9),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),X11))|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|~epred1_3(X9,X8,X7))&((((((~r1_funct_2(u1_struct_0(k1_tsep_1(X7,X9,X10)),u1_struct_0(X8),u1_struct_0(X9),u1_struct_0(X8),k10_tmap_1(X7,X8,X9,X10,X11,X12),X11)|m1_pre_topc(X10,X9))|~r2_funct_2(u1_struct_0(k2_tsep_1(X7,X9,X10)),u1_struct_0(X8),k3_tmap_1(X7,X8,X9,k2_tsep_1(X7,X9,X10),X11),k3_tmap_1(X7,X8,X10,k2_tsep_1(X7,X9,X10),X12)))|(~v1_funct_1(X12)|(~v1_funct_2(X12,u1_struct_0(X10),u1_struct_0(X8))|~m1_subset_1(X12,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X10),u1_struct_0(X8)))))))|(~v1_funct_1(X11)|(~v1_funct_2(X11,u1_struct_0(X9),u1_struct_0(X8))|~m1_subset_1(X11,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X9),u1_struct_0(X8)))))))|(v2_struct_0(X10)|~m1_pre_topc(X10,X7)))|~epred1_3(X9,X8,X7)))))), inference(distribute,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(shift_quantors,[status(thm)],[inference(variable_rename,[status(thm)],[inference(fof_nnf,[status(thm)],[inference(fof_simplification,[status(thm)],[c_0_86])])])])])])])).
cnf(c_0_145,negated_conjecture,(epred1_3(esk3_0,X1,esk1_0)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_42]), c_0_44]), c_0_45])]), c_0_46]), c_0_22]), ['final']).
cnf(c_0_146,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_tsep_1(X3,X4)|v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v5_pre_topc(X5,X3,X2)|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v5_pre_topc(X6,X4,X2)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r4_tsep_1(X1,X3,X4)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))), inference(split_conjunct,[status(thm)],[c_0_88])).
cnf(c_0_147,negated_conjecture,(~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_89, c_0_61])])).
cnf(c_0_148,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_149,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_150,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_151,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_152,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_153,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_154,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_155,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_156,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_157,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_158,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_159,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_160,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),X2,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_161,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),X2,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_72]), c_0_73]), c_0_61]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_162,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_163,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_164,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_165,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),X2,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_166,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_167,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_168,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_169,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_170,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_171,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_172,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_173,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_174,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),X3,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_175,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),X2,k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_176,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_177,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_90, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_178,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_179,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_91, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_180,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_181,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),X3,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_92, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_182,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_183,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),X3,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_93, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_184,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,esk3_0),X3),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,X2,esk3_0),esk5_0))|r1_tsep_1(X2,esk3_0)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0)),esk5_0)|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),X2,k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0)),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_18]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_185,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,X2,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,X2,k2_tsep_1(X1,X2,esk4_0),X3),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,X2,esk4_0),esk6_0))|r1_tsep_1(X2,esk4_0)|v2_struct_0(X1)|v2_struct_0(X2)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0)),esk6_0)|~r2_funct_2(u1_struct_0(X2),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),X2,k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0)),X3)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_94, c_0_24]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_186,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,esk3_0),X3,k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(X2,esk3_0)|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_187,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,X1),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,X1,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,X1),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,X1,esk5_0,X2),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,X1,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(X1,esk4_0)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_97, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_188,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_99, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_189,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_100, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_190,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_101, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_191,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_102, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_192,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_103, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_193,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_104, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_194,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_105, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_195,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_106, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_196,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_107, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_197,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_108, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_198,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,X1,esk3_0,X2,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(X1,esk4_0)|~r1_tsep_1(X1,esk3_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_109, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_199,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_110, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_200,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_111, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_201,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_202,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,X1))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X1,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X2),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,X1),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,X1,esk6_0,X2)))|v2_struct_0(X1)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(esk4_0,X1)|~r1_tsep_1(esk3_0,X1)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_113, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_203,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_72]), c_0_73]), c_0_61]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_204,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_114, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_205,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_206,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_207,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_115, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_208,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk3_0,esk3_0),X3,k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk3_0,esk3_0)|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_95, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_209,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_116, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_210,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X2,esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,X2,esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(X1,esk4_0,esk4_0),X3,k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(esk4_0,esk4_0)|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_96, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_211,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r4_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~v5_pre_topc(X3,X2,esk2_0)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_212,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_118, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_213,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r4_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~v5_pre_topc(X3,X2,esk2_0)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_214,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),X3,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|v2_struct_0(X2)|~r4_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~v5_pre_topc(X3,X2,esk2_0)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_215,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_119, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_216,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_120, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_217,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_121, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_218,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_122, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_219,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_123, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_220,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_124, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_221,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_125, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_222,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_126, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_223,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_127, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_224,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_128, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_225,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_129, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_226,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk5_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk5_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_130, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_227,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_131, c_0_66]), c_0_67]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_228,negated_conjecture,(~r1_tsep_1(esk3_0,esk4_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_229,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_132, c_0_74]), c_0_75]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_230,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,X1,esk3_0,X2,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,X1,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,X1,k2_tsep_1(esk1_0,X1,esk4_0),X2),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,X1,esk4_0),esk6_0))|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,X1,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,X1,k2_tsep_1(esk1_0,X1,esk3_0),X2),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,X1,esk3_0),esk5_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_133, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_231,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,X2,esk5_0,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_18]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_232,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,X2,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,X2)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,X2,esk6_0,X3))|v2_struct_0(X1)|v2_struct_0(X2)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_112, c_0_24]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_233,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_134, c_0_68]), c_0_69]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_234,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,X2,esk3_0,X3,esk5_0),k1_tsep_1(X1,X2,esk3_0),esk2_0)|v2_struct_0(X1)|v2_struct_0(X2)|~r4_tsep_1(X1,X2,esk3_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v5_pre_topc(X3,X2,esk2_0)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_18]), c_0_135]), c_0_19]), c_0_13]), c_0_20]), c_0_21])]), c_0_22]), c_0_23]), ['final']).
cnf(c_0_235,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,X2,esk4_0,X3,esk6_0),k1_tsep_1(X1,X2,esk4_0),esk2_0)|v2_struct_0(X1)|v2_struct_0(X2)|~r4_tsep_1(X1,X2,esk4_0)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v5_pre_topc(X3,X2,esk2_0)|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~r1_tsep_1(X2,esk4_0)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_117, c_0_24]), c_0_136]), c_0_25]), c_0_14]), c_0_20]), c_0_21])]), c_0_26]), c_0_23]), ['final']).
cnf(c_0_236,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X2,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(esk6_0,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(esk6_0,u1_struct_0(X2),u1_struct_0(esk2_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_137, c_0_72]), c_0_73]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_237,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tsep_1(esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_66]), c_0_67]), c_0_62]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_238,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X7=k10_tmap_1(X1,X2,X3,X4,X5,X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r1_tsep_1(X3,X4)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|~r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6)|~r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_239,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0)|~r1_tsep_1(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_138, c_0_18]), c_0_19]), c_0_13])]), ['final']).
cnf(c_0_240,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),X3),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_241,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X3),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_72]), c_0_73]), c_0_61]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_242,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),X3),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_243,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),X3),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_28, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_244,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_66]), c_0_67]), c_0_62]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_245,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_72]), c_0_73]), c_0_61]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_246,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_74]), c_0_75]), c_0_63]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_247,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),X3))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X2)|v2_struct_0(X1)|~m1_subset_1(X3,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X2),u1_struct_0(esk2_0))))|~v1_funct_2(X3,u1_struct_0(X2),u1_struct_0(esk2_0))|~v1_funct_1(X3)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(X2,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_17, c_0_68]), c_0_69]), c_0_64]), c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_248,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,X1,esk3_0,X2,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,X1,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,X1,k2_tsep_1(esk1_0,X1,esk4_0),X2),k3_tmap_1(esk1_0,esk2_0,esk4_0,k2_tsep_1(esk1_0,X1,esk4_0),esk6_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(X1,esk3_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_139, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_249,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,X1),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,X1,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,X1),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,X1,esk5_0,X2),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,X1,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,X1,esk4_0,X2,esk6_0)))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,X1)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,X1),esk5_0),k3_tmap_1(esk1_0,esk2_0,X1,k2_tsep_1(esk1_0,esk3_0,X1),X2))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(X1,esk4_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_140, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_250,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|v2_struct_0(X5)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,X3,X4),X5)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X3,k1_tsep_1(X1,X4,X5))),u1_struct_0(X2),k10_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X5,k10_tmap_1(X1,X2,X3,X4,X6,X7),X8),k10_tmap_1(X1,X2,X3,k1_tsep_1(X1,X4,X5),X6,k10_tmap_1(X1,X2,X4,X5,X7,X8)))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X5,X1)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X7)|~m1_subset_1(X8,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X5),u1_struct_0(X2))))|~v1_funct_2(X8,u1_struct_0(X5),u1_struct_0(X2))|~v1_funct_1(X8)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X6),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X7))|~r1_tsep_1(X3,X5)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X4,X5)),u1_struct_0(X2),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X4,X5),X7),k3_tmap_1(X1,X2,X5,k2_tsep_1(X1,X4,X5),X8))), inference(split_conjunct,[status(thm)],[c_0_56]), ['final']).
cnf(c_0_251,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,X1,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_141, c_0_18]), c_0_19]), c_0_13])]), ['final']).
cnf(c_0_252,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,X1,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_142, c_0_18]), c_0_19]), c_0_13])])).
cnf(c_0_253,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,X1,X2)|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_84, c_0_68]), c_0_69]), c_0_64]), c_0_43]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_254,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0)|~r1_tsep_1(esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_143, c_0_24]), c_0_25]), c_0_14])]), ['final']).
cnf(c_0_255,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|X7=k10_tmap_1(X1,X2,X3,X4,X5,X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|~r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6)|~r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_256,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,X7),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|X7!=k10_tmap_1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_257,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,X7),X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))|~m1_subset_1(X7,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))))|~v1_funct_2(X7,u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2))|~v1_funct_1(X7)|X7!=k10_tmap_1(X1,X2,X3,X4,X5,X6)), inference(split_conjunct,[status(thm)],[c_0_50]), ['final']).
cnf(c_0_258,plain,(v2_struct_0(X4)|m1_pre_topc(X1,X4)|~epred1_3(X1,X2,X3)|~m1_pre_topc(X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X5),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X6))|~r1_funct_2(u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X3,X2,X1,X4,X5,X6),X6)), inference(split_conjunct,[status(thm)],[c_0_144]), ['final']).
cnf(c_0_259,negated_conjecture,(epred1_3(esk3_0,esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_260,plain,(v2_struct_0(X4)|m1_pre_topc(X4,X1)|~epred1_3(X1,X2,X3)|~m1_pre_topc(X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X5),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X6))|~r1_funct_2(u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2),u1_struct_0(X1),u1_struct_0(X2),k10_tmap_1(X3,X2,X1,X4,X5,X6),X5)), inference(split_conjunct,[status(thm)],[c_0_144]), ['final']).
cnf(c_0_261,plain,(v2_struct_0(X4)|r1_funct_2(u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2),u1_struct_0(X1),u1_struct_0(X2),k10_tmap_1(X3,X2,X1,X4,X5,X6),X5)|~epred1_3(X1,X2,X3)|~m1_pre_topc(X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X5),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X6))|~m1_pre_topc(X4,X1)), inference(split_conjunct,[status(thm)],[c_0_144]), ['final']).
cnf(c_0_262,plain,(v2_struct_0(X4)|r1_funct_2(u1_struct_0(k1_tsep_1(X3,X1,X4)),u1_struct_0(X2),u1_struct_0(X4),u1_struct_0(X2),k10_tmap_1(X3,X2,X1,X4,X5,X6),X6)|~epred1_3(X1,X2,X3)|~m1_pre_topc(X4,X3)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X1),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X3,X1,X4)),u1_struct_0(X2),k3_tmap_1(X3,X2,X1,k2_tsep_1(X3,X1,X4),X5),k3_tmap_1(X3,X2,X4,k2_tsep_1(X3,X1,X4),X6))|~m1_pre_topc(X1,X4)), inference(split_conjunct,[status(thm)],[c_0_144]), ['final']).
cnf(c_0_263,plain,(v5_pre_topc(k10_tmap_1(X1,X2,X3,X4,X5,X6),k1_tsep_1(X1,X3,X4),X2)|v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|~r4_tsep_1(X1,X3,X4)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v5_pre_topc(X6,X4,X2)|~v5_pre_topc(X5,X3,X2)|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X6)|~v1_funct_1(X5)|~m1_pre_topc(X4,X1)|~m1_pre_topc(X3,X1)|~l1_pre_topc(X2)|~l1_pre_topc(X1)|~v2_pre_topc(X2)|~v2_pre_topc(X1)), inference(csr,[status(thm)],[c_0_146, c_0_117]), ['final']).
cnf(c_0_264,negated_conjecture,(~m1_subset_1(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_147, c_0_73])])).
cnf(c_0_265,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_funct_2(u1_struct_0(k1_tsep_1(X1,X3,X4)),u1_struct_0(X2),u1_struct_0(k1_tsep_1(X1,X4,X3)),u1_struct_0(X2),k10_tmap_1(X1,X2,X3,X4,X5,X6),k10_tmap_1(X1,X2,X4,X3,X6,X5))|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))), inference(split_conjunct,[status(thm)],[c_0_80]), ['final']).
cnf(c_0_266,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(X3),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X3,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X5)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))), inference(split_conjunct,[status(thm)],[c_0_76]), ['final']).
cnf(c_0_267,plain,(v2_struct_0(X1)|v2_struct_0(X2)|v2_struct_0(X3)|v2_struct_0(X4)|r1_tsep_1(X3,X4)|r2_funct_2(u1_struct_0(X4),u1_struct_0(X2),k3_tmap_1(X1,X2,k1_tsep_1(X1,X3,X4),X4,k10_tmap_1(X1,X2,X3,X4,X5,X6)),X6)|~l1_pre_topc(X1)|~v2_pre_topc(X1)|~l1_pre_topc(X2)|~v2_pre_topc(X2)|~m1_pre_topc(X3,X1)|~m1_pre_topc(X4,X1)|~m1_subset_1(X5,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X3),u1_struct_0(X2))))|~v1_funct_2(X5,u1_struct_0(X3),u1_struct_0(X2))|~v1_funct_1(X5)|~m1_subset_1(X6,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X4),u1_struct_0(X2))))|~v1_funct_2(X6,u1_struct_0(X4),u1_struct_0(X2))|~v1_funct_1(X6)|~r2_funct_2(u1_struct_0(k2_tsep_1(X1,X3,X4)),u1_struct_0(X2),k3_tmap_1(X1,X2,X3,k2_tsep_1(X1,X3,X4),X5),k3_tmap_1(X1,X2,X4,k2_tsep_1(X1,X3,X4),X6))), inference(split_conjunct,[status(thm)],[c_0_76]), ['final']).
cnf(c_0_268,negated_conjecture,(epred1_3(esk4_0,X1,esk1_0)|v2_struct_0(X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_87, c_0_43]), c_0_44]), c_0_45])]), c_0_46]), c_0_26]), ['final']).
cnf(c_0_269,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_270,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_271,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_272,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_273,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_274,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_275,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_276,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_277,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_278,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_279,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_280,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_281,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_282,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_283,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_284,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_285,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_286,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_287,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_288,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_289,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_290,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_291,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_292,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_293,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_294,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_295,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_296,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_297,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_298,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_299,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_300,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_301,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_302,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_303,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_304,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_305,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_306,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_307,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_148, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_308,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_309,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_310,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_149, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_311,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_312,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_313,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_314,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_150, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_315,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_316,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_317,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_151, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_318,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_319,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_320,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_321,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_322,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_323,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_324,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_325,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_152, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_326,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_327,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_328,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_329,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_330,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_331,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_332,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_333,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_153, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_334,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_335,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_336,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_337,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_338,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_339,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_154, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_340,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_341,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_155, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_342,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_343,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_344,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_345,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_346,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_347,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_348,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_156, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_349,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_350,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_351,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_352,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_353,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_354,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_355,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_356,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_162, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_357,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_358,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_359,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_360,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_163, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_361,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_362,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_363,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_164, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_364,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_365,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_366,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_367,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_368,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_157, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_369,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_370,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_158, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_371,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_372,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_373,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_159, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_374,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_375,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_376,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_377,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_378,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_379,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_380,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_381,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_382,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_383,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_384,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_385,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_386,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_387,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_166, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_388,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_389,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_390,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_391,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_392,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_393,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_394,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_395,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_396,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_397,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_398,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_399,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_400,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_401,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_402,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_403,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_404,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_405,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_160, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_406,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_407,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_408,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_409,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_410,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_411,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_412,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_413,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_414,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_161, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_415,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_416,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_417,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_418,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_419,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_420,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_421,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_422,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_423,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_165, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_424,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_425,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_426,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_427,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_428,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_429,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_430,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_431,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_432,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_433,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_434,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_435,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_436,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_437,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_438,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_439,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_167, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_440,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_175, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_441,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_442,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_443,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_444,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_445,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_168, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_446,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_176, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_447,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_448,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_449,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_177, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_450,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_451,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_178, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_452,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk5_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0)),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_453,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_454,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_455,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_179, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_456,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_457,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_180, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_458,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk6_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_459,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_460,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_181, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_461,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_462,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_463,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_182, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_464,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_465,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_466,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_467,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_468,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_469,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_183, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_470,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_471,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_472,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_473,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_169, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_474,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_475,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_476,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_477,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_478,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_170, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_479,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_171, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_480,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_172, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_481,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk5_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_173, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_482,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_483,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk6_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_174, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_484,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_188, c_0_66]), c_0_67])]), ['final']).
cnf(c_0_485,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_189, c_0_74]), c_0_75])]), ['final']).
cnf(c_0_486,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk5_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_487,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_190, c_0_68]), c_0_69])]), ['final']).
cnf(c_0_488,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_191, c_0_72]), c_0_73])]), ['final']).
cnf(c_0_489,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_192, c_0_66]), c_0_67])]), ['final']).
cnf(c_0_490,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk6_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_491,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_193, c_0_74]), c_0_75])]), ['final']).
cnf(c_0_492,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_194, c_0_68]), c_0_69])]), ['final']).
cnf(c_0_493,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_195, c_0_72]), c_0_73])]), ['final']).
cnf(c_0_494,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_196, c_0_66]), c_0_67])]), ['final']).
cnf(c_0_495,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_197, c_0_74]), c_0_75])]), ['final']).
cnf(c_0_496,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_497,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_199, c_0_68]), c_0_69])]), ['final']).
cnf(c_0_498,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_200, c_0_72]), c_0_73])]), ['final']).
cnf(c_0_499,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_500,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_501,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_502,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_503,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_504,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_505,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_506,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_204, c_0_66]), c_0_67])]), ['final']).
cnf(c_0_507,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_508,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_509,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_510,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_511,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_512,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_207, c_0_74]), c_0_75])]), ['final']).
cnf(c_0_513,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk5_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0)),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_514,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_515,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_516,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_517,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_518,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_519,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_209, c_0_68]), c_0_69])]), ['final']).
cnf(c_0_520,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk6_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_521,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(X1,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_522,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_523,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_524,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(X1,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_525,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(X1,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_526,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_212, c_0_72]), c_0_73])]), ['final']).
cnf(c_0_527,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_528,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_529,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(X1,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_530,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_531,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_532,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_201, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_533,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_534,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_535,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_203, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_536,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_186, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_537,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_538,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_539,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_540,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_541,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_205, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_542,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_543,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_544,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_18]), c_0_135]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_545,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_211, c_0_24]), c_0_136]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_546,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_547,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk5_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0)),esk5_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_548,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_18]), c_0_135]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_549,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_213, c_0_24]), c_0_136]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_550,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk3_0,esk3_0),esk6_0,k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,esk3_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_551,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(X1,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(X1,esk3_0,esk3_0),esk5_0,k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_208, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_552,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(X1,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(X1,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(X1,esk4_0,esk4_0),esk6_0,k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_210, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_553,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk6_0))|r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)),esk6_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_554,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_555,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_556,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_187, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_557,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_215, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_558,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_216, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_559,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_217, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_560,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_218, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_561,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_219, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_562,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_220, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_563,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_564,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_206, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_565,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_221, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_566,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_567,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_18]), c_0_135]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_568,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_214, c_0_24]), c_0_136]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_569,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_222, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_570,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_223, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_571,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_224, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_572,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_225, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_573,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_574,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0)),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_575,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_576,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_226, c_0_18]), c_0_19])]), c_0_22]), ['final']).
cnf(c_0_577,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_227, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_578,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_579,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0)),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_580,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0)),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~r1_tsep_1(esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk1_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_581,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,esk3_0),esk6_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk4_0,esk3_0),esk5_0))|r1_tsep_1(esk4_0,esk3_0)|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk3_0),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_582,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,esk4_0),esk5_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk3_0,esk4_0),esk6_0))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk4_0),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)|~m1_pre_topc(esk4_0,X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_18]), c_0_19]), c_0_13])]), c_0_228]), c_0_22]), ['final']).
cnf(c_0_583,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_229, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_584,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0))), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_230, c_0_98]), c_0_18]), c_0_19]), c_0_13]), c_0_42])]), c_0_22]), ['final']).
cnf(c_0_585,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_586,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_587,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_588,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_589,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_590,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_233, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_591,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_592,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_593,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_594,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_595,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_596,negated_conjecture,(m1_subset_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_236, c_0_24]), c_0_25])]), c_0_26]), ['final']).
cnf(c_0_597,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,X1,esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk3_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_237, c_0_18]), c_0_19]), c_0_13])]), ['final']).
cnf(c_0_598,negated_conjecture,(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,X1,esk5_0)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk3_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk3_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_239]), c_0_66]), c_0_18]), c_0_67]), c_0_19]), c_0_62]), c_0_13]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_599,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_600,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_601,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_602,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_603,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_604,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_240, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_605,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_606,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_607,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~r1_tsep_1(esk4_0,esk4_0)|~r1_tsep_1(esk4_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_198, c_0_24]), c_0_25]), c_0_14]), c_0_43])]), c_0_26]), ['final']).
cnf(c_0_608,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_609,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_610,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_611,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_612,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_613,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_614,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_244, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_615,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_616,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_617,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_618,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_619,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_620,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_621,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_66]), c_0_67]), c_0_62])]), ['final']).
cnf(c_0_622,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|~r1_tsep_1(esk4_0,esk3_0)|~r1_tsep_1(esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_202, c_0_18]), c_0_19]), c_0_13]), c_0_42])]), c_0_22]), ['final']).
cnf(c_0_623,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_624,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_625,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_626,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_627,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_241, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_628,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|~r1_tsep_1(esk3_0,esk3_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_248, c_0_98]), c_0_18]), c_0_19]), c_0_13]), c_0_42])]), c_0_22]), ['final']).
cnf(c_0_629,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_630,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_631,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_632,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_633,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_634,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_635,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_245, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_636,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_637,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_638,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_639,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_640,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_641,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk3_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_72]), c_0_73]), c_0_61])]), ['final']).
cnf(c_0_642,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|~r1_tsep_1(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_249, c_0_98]), c_0_24]), c_0_25]), c_0_14]), c_0_43])]), c_0_26]), ['final']).
cnf(c_0_643,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_644,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_645,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0))),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,X1,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,X1,esk3_0,X2,esk5_0),esk6_0),k10_tmap_1(esk1_0,esk2_0,X1,k1_tsep_1(esk1_0,esk3_0,esk4_0),X2,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)))|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(k2_tsep_1(esk1_0,X1,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,X1,k2_tsep_1(esk1_0,X1,esk3_0),X2),k3_tmap_1(esk1_0,esk2_0,esk3_0,k2_tsep_1(esk1_0,X1,esk3_0),esk5_0))|~m1_subset_1(X2,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(X1),u1_struct_0(esk2_0))))|~v1_funct_2(X2,u1_struct_0(X1),u1_struct_0(esk2_0))|~v1_funct_1(X2)|~r1_tsep_1(X1,esk4_0)|~m1_pre_topc(X1,esk1_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_250, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_646,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk6_0)|~r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_251, c_0_24]), c_0_25]), c_0_14])]), ['final']).
cnf(c_0_647,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_648,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_649,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_650,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_242, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_651,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),esk5_0)|~r1_tsep_1(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_252, c_0_24]), c_0_25]), c_0_14])]), ['final']).
cnf(c_0_652,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_653,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_654,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_655,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_656,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_246, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_657,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk3_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_658,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_659,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_660,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_661,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk3_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk3_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_74]), c_0_75]), c_0_63])]), ['final']).
cnf(c_0_662,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_663,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0),esk2_0)|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~r4_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),esk2_0)|~r1_tsep_1(k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_664,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),X1)|k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)!=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,X1,esk6_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_253, c_0_24]), c_0_25]), c_0_14])]), ['final']).
cnf(c_0_665,negated_conjecture,(k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)=k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,X1,esk6_0)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),X1)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(esk4_0),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(esk4_0),u1_struct_0(esk2_0))|~v1_funct_1(X1)|~r1_tsep_1(esk4_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_238, c_0_254]), c_0_68]), c_0_24]), c_0_69]), c_0_25]), c_0_64]), c_0_14]), c_0_43]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_666,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_667,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_668,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0),u1_struct_0(k1_tsep_1(X1,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0)),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_243, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_669,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk3_0,esk3_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,esk3_0),esk5_0),k3_tmap_1(X1,esk2_0,esk3_0,k2_tsep_1(X1,esk3_0,esk3_0),esk5_0))|r1_tsep_1(esk3_0,esk3_0)|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk3_0,esk3_0),esk3_0,k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0)),esk5_0)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_184, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_670,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),k1_tsep_1(esk1_0,esk4_0,esk4_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_671,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk5_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_672,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),esk6_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_247, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_673,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_35, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_674,negated_conjecture,(v1_funct_2(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),u1_struct_0(k1_tsep_1(X1,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0))),u1_struct_0(esk2_0))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_37, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_675,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk3_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk5_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_32, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_676,negated_conjecture,(v1_funct_1(k10_tmap_1(X1,esk2_0,esk4_0,k1_tsep_1(esk1_0,esk4_0,esk4_0),esk6_0,k10_tmap_1(esk1_0,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)))|v2_struct_0(k1_tsep_1(esk1_0,esk4_0,esk4_0))|v2_struct_0(X1)|~m1_pre_topc(k1_tsep_1(esk1_0,esk4_0,esk4_0),X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_33, c_0_68]), c_0_69]), c_0_64])]), ['final']).
cnf(c_0_677,negated_conjecture,(r2_funct_2(u1_struct_0(k2_tsep_1(X1,esk4_0,esk4_0)),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,esk4_0),esk6_0),k3_tmap_1(X1,esk2_0,esk4_0,k2_tsep_1(X1,esk4_0,esk4_0),esk6_0))|r1_tsep_1(esk4_0,esk4_0)|v2_struct_0(X1)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(X1,esk2_0,k1_tsep_1(X1,esk4_0,esk4_0),esk4_0,k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0)),esk6_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_185, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_678,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_679,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk3_0,esk3_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk3_0,esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0))|v2_struct_0(X1)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_231, c_0_18]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_680,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(X1,esk4_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(X1,esk4_0,esk4_0)),u1_struct_0(esk2_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0))|v2_struct_0(X1)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_232, c_0_24]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_681,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0),k1_tsep_1(X1,esk4_0,esk3_0),esk2_0)|v2_struct_0(X1)|~r4_tsep_1(X1,esk4_0,esk3_0)|~r1_tsep_1(esk4_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_24]), c_0_136]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_682,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk3_0,esk3_0,esk5_0,esk5_0),k1_tsep_1(X1,esk3_0,esk3_0),esk2_0)|v2_struct_0(X1)|~r4_tsep_1(X1,esk3_0,esk3_0)|~r1_tsep_1(esk3_0,esk3_0)|~m1_pre_topc(esk3_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_234, c_0_18]), c_0_135]), c_0_19]), c_0_13])]), c_0_22]), ['final']).
cnf(c_0_683,negated_conjecture,(v5_pre_topc(k10_tmap_1(X1,esk2_0,esk4_0,esk4_0,esk6_0,esk6_0),k1_tsep_1(X1,esk4_0,esk4_0),esk2_0)|v2_struct_0(X1)|~r4_tsep_1(X1,esk4_0,esk4_0)|~r1_tsep_1(esk4_0,esk4_0)|~m1_pre_topc(esk4_0,X1)|~l1_pre_topc(X1)|~v2_pre_topc(X1)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_235, c_0_24]), c_0_136]), c_0_25]), c_0_14])]), c_0_26]), ['final']).
cnf(c_0_684,negated_conjecture,(X1=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)|~r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,X1),esk6_0)|~r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,X1),esk5_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_255, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_685,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,X1),esk5_0)|X1!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_256, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_686,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,X1),esk6_0)|X1!=k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)|~m1_subset_1(X1,k1_zfmisc_1(k2_zfmisc_1(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))))|~v1_funct_2(X1,u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0))|~v1_funct_1(X1)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_257, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_687,negated_conjecture,(m1_pre_topc(esk3_0,esk4_0)|~r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_258, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43])]), c_0_26]), c_0_259])]), ['final']).
cnf(c_0_688,negated_conjecture,(m1_pre_topc(esk4_0,esk3_0)|~r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_260, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43])]), c_0_26]), c_0_259])]), ['final']).
cnf(c_0_689,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(esk3_0),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk5_0)|~m1_pre_topc(esk4_0,esk3_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_261, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43])]), c_0_26]), c_0_259])]), ['final']).
cnf(c_0_690,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(esk4_0),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),esk6_0)|~m1_pre_topc(esk3_0,esk4_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_262, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43])]), c_0_26]), c_0_259])]), ['final']).
cnf(c_0_691,negated_conjecture,(v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)|~r4_tsep_1(esk1_0,esk3_0,esk4_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_263, c_0_98]), c_0_24]), c_0_18]), c_0_136]), c_0_135]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_692,negated_conjecture,(~v5_pre_topc(k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k1_tsep_1(esk1_0,esk3_0,esk4_0),esk2_0)), inference(cn,[status(thm)],[inference(rw,[status(thm)],[c_0_264, c_0_72])]), ['final']).
cnf(c_0_693,negated_conjecture,(r1_funct_2(u1_struct_0(k1_tsep_1(esk1_0,esk3_0,esk4_0)),u1_struct_0(esk2_0),u1_struct_0(k1_tsep_1(esk1_0,esk4_0,esk3_0)),u1_struct_0(esk2_0),k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0),k10_tmap_1(esk1_0,esk2_0,esk4_0,esk3_0,esk6_0,esk5_0))), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_265, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_694,negated_conjecture,(r2_funct_2(u1_struct_0(esk3_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk3_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk5_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_266, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_228]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_695,negated_conjecture,(r2_funct_2(u1_struct_0(esk4_0),u1_struct_0(esk2_0),k3_tmap_1(esk1_0,esk2_0,k1_tsep_1(esk1_0,esk3_0,esk4_0),esk4_0,k10_tmap_1(esk1_0,esk2_0,esk3_0,esk4_0,esk5_0,esk6_0)),esk6_0)), inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_267, c_0_98]), c_0_24]), c_0_18]), c_0_25]), c_0_19]), c_0_14]), c_0_13]), c_0_43]), c_0_42]), c_0_20]), c_0_44]), c_0_21]), c_0_45])]), c_0_228]), c_0_26]), c_0_22]), c_0_23]), c_0_46]), ['final']).
cnf(c_0_696,negated_conjecture,(epred1_3(esk3_0,esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_145, c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_697,negated_conjecture,(epred1_3(esk4_0,esk1_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_44]), c_0_45])]), c_0_46]), ['final']).
cnf(c_0_698,negated_conjecture,(epred1_3(esk4_0,esk2_0,esk1_0)), inference(sr,[status(thm)],[inference(cn,[status(thm)],[inference(rw,[status(thm)],[inference(spm,[status(thm)],[c_0_268, c_0_20]), c_0_21])]), c_0_23]), ['final']).
cnf(c_0_699,negated_conjecture,(v1_borsuk_1(esk4_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
cnf(c_0_700,negated_conjecture,(v1_borsuk_1(esk3_0,esk1_0)), inference(split_conjunct,[status(thm)],[c_0_11]), ['final']).
# SZS output end Saturation
# Proof object total steps             : 701
# Proof object clause steps            : 678
# Proof object formula steps           : 23
# Proof object conjectures             : 651
# Proof object clause conjectures      : 648
# Proof object formula conjectures     : 3
# Proof object initial clauses used    : 52
# Proof object initial formulas used   : 9
# Proof object generating inferences   : 622
# Proof object simplifying inferences  : 2406
# Parsed axioms                        : 9
# Removed by relevancy pruning/SinE    : 0
# Initial clauses                      : 58
# Removed in clause preprocessing      : 0
# Initial clauses in saturation        : 58
# Processed clauses                    : 763
# ...of these trivial                  : 0
# ...subsumed                          : 34
# ...remaining for further processing  : 729
# Other redundant clauses eliminated   : 0
# Clauses deleted for lack of memory   : 0
# Backward-subsumed                    : 3
# Backward-rewritten                   : 3
# Generated clauses                    : 657
# ...of the previous two non-trivial   : 653
# Contextual simplify-reflections      : 32
# Paramodulations                      : 657
# Factorizations                       : 0
# Equation resolutions                 : 0
# Current number of processed clauses  : 671
#    Positive orientable unit clauses  : 36
#    Positive unorientable unit clauses: 0
#    Negative unit clauses             : 6
#    Non-unit-clauses                  : 629
# Current number of unprocessed clauses: 0
# ...number of literals in the above   : 0
# Current number of archived formulas  : 0
# Current number of archived clauses   : 58
# Clause-clause subsumption calls (NU) : 505235
# Rec. Clause-clause subsumption calls : 118
# Non-unit clause-clause subsumptions  : 47
# Unit Clause-clause subsumption calls : 168
# Rewrite failures with RHS unbound    : 0
# BW rewrite match attempts            : 32
# BW rewrite match successes           : 3
# Condensation attempts                : 0
# Condensation successes               : 0
# Termbank termtop insertions          : 172948

# -------------------------------------------------
# User time                : 0.210 s
# System time              : 0.000 s
# Total time               : 0.210 s
# Maximum resident set size: 59184 pages
